## Homework 0
### Due Date: September 9 at 11:59PM ET

Before starting Homework 1, you should get practice with submission of assignments. Homework 0 has a few questions but you will not receive a grade for your answers. Focus on  

- fetching
- modifying
- validating
- submitting  

your Jupyter notebook. Consult the [instructional video](https://nbgrader.readthedocs.io/en/stable/user_guide/highlights.html#student-assignment-list-extension-for-jupyter-notebooks) for more information about the grading component of JupyterHub.

#### Collaboration Policy

You are prohibited from sharing your solutions. You are discouraged from providing guidance or assistance on questions. You should refrain from discussing the assignment with others. Consult the syllabus for information about policies surrounding misconduct.

**Name:** *Alina Wang*

**NetId:** *hw3414*

### Rubric

Question | Points
--- | ---
[Question 1](#q1) | 3
Total | 3

### Question 1: String Formatting <a name='q1'></a>

A string is a collection of characters meaning letters, numbers and punctuation. We can specify a string through quotation marks 

> `"Hello World"`

If we want to display a string to the output of a cell, then we can use the `print` function.



```python
print("Hello World")
```

Output:
```
Hello World
```

```python
us_states = ["s","a","d","h"]
print(us_states[::2])
```

Output:
```
['s', 'd']
```

Sometimes we want to add information to a string that depends on a calculation. We cannot include the operation between quotation marks. 

```python
print("The sum of 3 and 2 is 3+2")
```

Output:
```
The sum of 3 and 2 is 3+2
```

Note that the string does not appear as `"The sum of 3 and 2 is 5"` in the output of the cell.

> So how can we display the result of a calculation in a string? 

We need a way to insert the missing number. We will explore three ways to add a placeholder for a number in a string. 

#### Question 1.1 

We can use a percentage sign to indicate a placeholder.

```python
print("The sum of 3 and 2 is %d" % (2 + 3))
```

Here the placeholder is `%d`. Note that `d` indicates an integer number like `5`. Suppose we changed from `d` to `f`. 

```python
print("The sum of 3 and 2 is %f" % (2 + 3))
```

What is the output of the code?

- A. `The sum of 3 and 2 is 5`
- B. `The sum of 3 and 2 is 5.000000`
- C. `The sum of 3 and 2 is 2 + 3`

Indicate your answer as the string `"A"`,`"B"` or `"C"`. 

<!-- grade_id: cell-92e5b8cdb6739630, points: 0, grade: false, solution: true -->
```python
q1_1 = "B"

# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[15], line 4
      1 q1_1 = "C"
      3 # YOUR CODE HERE
----> 4 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-d703f20b90b0fa40, points: 1, grade: true, solution: false -->
```python
# TEST 

assert q1_1 in ["A","B","C"]

```

#### Question 1.2

We can use the `format` method to insert a number into a placeholder.

```python
print("The sum of 3 and 2 is {}".format(2+3))
```

Here the placeholder consists of braces `{}`. Note that `format` is an operation on the string that inserts the number into the placeholder.

**True or False** The following code 

```python 
print("The sum of 3 and 2 is {placeholder_name}".format(placeholder_name = 2 + 3))
```

generates an error message.

Indicate your answer as the string `"True"` or `"False"`.

```python
print("The sum of 3 and 2 is {placeholder_name}".format(placeholder_name = 2 + 3))
```

Output:
```
The sum of 3 and 2 is 5
```

<!-- grade_id: cell-8513e795e6987ddd, points: 0, grade: false, solution: true -->
```python
q1_2 = "True" 

# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[18], line 4
      1 q1_2 = "True" 
      3 # YOUR CODE HERE
----> 4 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-aec5dacf8a7ffedb, points: 1, grade: true, solution: false -->
```python
# TEST 

assert q1_2 in ["True", "False"]

```

#### Question 1.3

We can append the character `f` in front of the string.

```python
print(f"The sum of 3 and 2 is {2+3:f}")
```

Output:
```
The sum of 3 and 2 is 5.000000
```

Here the placeholder consists of braces `{}`. We need to enter the number between the braces. Suppose we made a modification. 

```python
print(f"The sum of 3 and 2 is {2+3:f}")
```
Does the code cause an error message to appear in the output of the cell?

Enter you answer as the string `"True"`,`"False"`.

```python
print(f"The sum of 3 and 2 is {2+3:d}")
```

Output:
```
The sum of 3 and 2 is 5
```

<!-- grade_id: cell-17e45add4e6a199f, points: 0, grade: false, solution: true -->
```python
q1_3 = "False" 

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-844eafff1b47e1d3, points: 1, grade: true, solution: false -->
```python
# TEST 

assert q1_3 in ["True", "False"]

```