## Homework 1 : Control Statements

#### Due Date: Saturday September 16 at 11:59PM ET

In this assignment, we will learn to organize code through

- if statements 
- for loops and while loops

Along the way, we will explore some of the tools from Week 2. The questions guide you step-by-step through these approaches. Do not post questions to Slack. 

#### Collaboration Policy

You are prohibited from sharing your solutions. You are discouraged from providing guidance or assistance on questions. You should refrain from discussing the assignment with others. Consult the syllabus for information about policies surrounding misconduct.

**Name:** *Alina Wang*

**NetId:** *hw3414*

### Rubric

Question | Points
--- | ---
Question 1.1 | 1
Question 1.2 | 1
Question 1.3 | 1
Question 1.4 | 1
Question 1.5 | 1
Question 1.6 | 1
Question 1.7 | 1
Question 2.1 | 1
Question 2.2 | 1
Question 2.3 | 1
Question 2.4 | 1
Total | 11

```python
from helper_functions import generate_expert_prediction
from helper_functions import get_consecutive
from helper_functions import plot_expert_forecasts
```

We will learn about functions and packages in Week 2

- A function allows us to label control statements just like a variable allows us to label data. 
- If we want to share functions and variables with collaborators we can make a package. We import the functions and variables from packages to take advantage of collaborators' code.

Please execute the above cell to import the functions we will need for Question 1 from the package `helper_functions`. 

<!-- grade_id: cell-74e9f6d9f93d6bc6, points: 0, grade: true, solution: false -->
```python
# TEST 

assert "generate_expert_prediction" in locals()
assert "plot_expert_forecasts" in locals()
assert "get_consecutive" in locals()
```

Note that you will be responsible for importing functions from the package `helper_functions` in Question 2.

If you managed to import the functions we will need for Question 1 from the package `helper_functions`, then the above cell will not generate an error message. 

### 1. Forecasting  

Companies need to determine forecasts to manage their operations. If experts can provide forecasts, then companies can inform their decisions from the advice of experts. Here we will study three experts 

- Expert `A`
- Expert `B`
- Expert `C`

Each expert can make a correct forecast or incorrect forecast. We want to determine the chance that the expert will make the correct prediction. We will assess the forecasts on 

- average number of correct forecast
- consistency of correct forecasts

Here consistency means the number of consecutive correction forecasts. The experts need to make forecasts in different market conditions. We will capture market conditions with numbers such as 1,2,3,4,....



#### Question 1.1

We want to allow experts `A`,`B`,`C` to make `100` forecasts in market conditions `10`. 

- Assign the integer number `100` to the variable `number_forecasts`
- Assign the integer number `10` to the variable `market_condition`
- Respectively assign the strings `A`,`B`,`C` to variables `expert_A`,`expert_B`,`expert_C`

<!-- grade_id: cell-584e947774ffa149, points: 0, grade: false, solution: true -->
```python
number_forecasts = 100
market_condition = 10
expert_A = 'A'
expert_B = 'B'
expert_C = 'C'

# YOUR CODE HERE
```

Execute the following three cells to generate line-charts.

```python
plot_expert_forecasts(expert_A, number_forecasts, market_condition)
```

```python
plot_expert_forecasts(expert_B, number_forecasts, market_condition)
```

```python
plot_expert_forecasts(expert_C, number_forecasts, market_condition)
```

Which expert made the largest number of correct forecasts? Assign the string `A`, `B` or `C` to the variable `q11`.

<!-- grade_id: cell-f6606f7a8ae2d2b6, points: 0, grade: false, solution: true -->
```python
q11 = 'A'

# YOUR CODE HERE
```

<!-- grade_id: cell-3173068f98658409, points: 1, grade: true, solution: false -->
```python
# TEST 

assert q11 in ["A","B","C"]

```

#### Question 1.2

We want to determine the average number of correct forecasts of experts `A`,`B`,`C`. 

Suppose we want to determine the average number of correct forecasts of expert `A`. We will allow expert `A` to make 100 forecasts in market conditions 1. 

```python
expert = "A"
number_forecasts = 100
market_conditions = 1

forecasts = generate_expert_prediction(expert, number_forecasts, market_conditions)
print(forecasts)
```

Here we use a function `generate_expert_prediction` from the package `helper_functions` to obtain 100 forecasts from expert `A` in market conditions 1. Note that the variable `forecasts` is a list containing 

- `1.` for correct forecast
- `0.` for incorrect forecast

If we want to count the number of correct forecasts, then we need to count the number of entries with `1.`. 

```python
number_correct = 0 
for forecast in forecasts:
    if (forecast == 1):
        number_correct = number_correct + 1
```

Here we use both an `if`-statement and `for`-statement. The variable `number_correct` contains the number of entries with `1.` in the list `forecasts`.

Using the value of the variables `number_correct` and `number_forecasts` determine the fraction of correct forecasts by expert `A`. Your answer must be a floating point number correct to three decimal places.

<!-- grade_id: cell-6d341117143a840e, points: 0, grade: false, solution: true -->
```python
fraction_correct_forecast_A = number_correct / number_forecasts
#print(fraction_correct_forecast_A)

# YOUR CODE HERE
```

<!-- grade_id: cell-cab182f8645252a7, points: 1, grade: true, solution: false -->
```python
# TEST 

assert 0 < fraction_correct_forecast_A < 1

```

Note that we could repeat the calculation for expert `B` and expert `C`.

#### Question 1.3

Some experts might have more accurate forecasts in some market conditions. We want to calculate the fraction of correct forecasts of experts `A`,`B`,`C` in different market condition.

Suppose we want to determine the fraction of correct forecasts of expert `B` in market conditions 

> 1,2,3,...,98,99,100

We will allow expert `B` to make 1000 forecasts in each of the 100 

```python
expert = "B"
number_forecasts = 1000
market_conditions = range(1,101)
```

We can use the code from Question 1.2 to compute the fraction of correct forecasts for particular market conditions.

```python
forecasts = generate_expert_prediction(expert, number_forecasts, market_conditions[0])

number_correct = 0 
for forecast in forecasts:
    if (forecast == 1):
        number_correct = number_correct + 1
        
print(number_correct / number_forecasts)
```

Here we learn that expert `B` made correct predictions 51.6\% of the time in market conditions `1`. 

Note that the market conditions are stored in the list `market_conditions`. So the entry with index `0` is the first entry of the list, namely, the integer number `1`. 

If we want to repeat the code for the other 99 market conditions, then we need to use another loop.

```python
fraction_correct = []
for market_condition in market_conditions:
    forecasts = generate_expert_prediction(expert, number_forecasts, market_condition)

    number_correct = 0 
    for forecast in forecasts:
        if (forecast == 1):
            number_correct = number_correct + 1
        
    fraction_correct.append(number_correct/number_forecasts)
    
print(fraction_correct)
```

We have two `for`-loops

- a loop through the 100 market conditions
- a loop through 1000 forecasts

Note that we store the fraction of correct forecasts in a list called `fraction_correct`. Compute the average of these fractions. Note that your answer must be a floating poit number correct to three decimal places. 

<!-- grade_id: cell-08f7c5b916ea3087, points: 0, grade: false, solution: true -->
```python
a = 0;
for fraction in fraction_correct:
    a = a + fraction;
average_fraction_correct_forecast_B = a / 100
print("%.3f" % average_fraction_correct_forecast_B)
# YOUR CODE HERE

```

<!-- grade_id: cell-a0f2e989ff993b2f, points: 1, grade: true, solution: false -->
```python
# TEST 

assert 0 < average_fraction_correct_forecast_B < 1

```

#### Question 1.4

Repeat the calculation from Question 1.3 for expert `A` and `C`. Based on the average fraction of correct predictions, we can rank the experts 

- `First`
- `Second`
- `Third`

Note that `First` corresponds to the largest average fraction of correct predictions and `Third` corresponds to the smallest average fraction of correct predictions. Fill in the keys of the dictionary `expert_averages` with the string 

- `A`
- `B`
- `C`

```python
expert = "A"
number_forecasts = 1000
market_conditions = range(1,101)
forecasts = generate_expert_prediction(expert, number_forecasts, market_conditions[0])

number_correct = 0 
for forecast in forecasts:
    if (forecast == 1):
        number_correct = number_correct + 1
        
print(number_correct / number_forecasts)
fraction_correct = []
for market_condition in market_conditions:
    forecasts = generate_expert_prediction(expert, number_forecasts, market_condition)

    number_correct = 0 
    for forecast in forecasts:
        if (forecast == 1):
            number_correct = number_correct + 1
        
    fraction_correct.append(number_correct/number_forecasts)
    
print(fraction_correct)
a = 0;
for fraction in fraction_correct:
    a = a + fraction;
average_fraction_correct_forecast_A = a / 100
print("%.3f" % average_fraction_correct_forecast_A)
# YOUR CODE HERE
```

```python
expert = "C"
number_forecasts = 1000
market_conditions = range(1,101)
forecasts = generate_expert_prediction(expert, number_forecasts, market_conditions[0])

number_correct = 0 
for forecast in forecasts:
    if (forecast == 1):
        number_correct = number_correct + 1
        
print(number_correct / number_forecasts)
fraction_correct = []
for market_condition in market_conditions:
    forecasts = generate_expert_prediction(expert, number_forecasts, market_condition)

    number_correct = 0 
    for forecast in forecasts:
        if (forecast == 1):
            number_correct = number_correct + 1
        
    fraction_correct.append(number_correct/number_forecasts)
    
print(fraction_correct)
a = 0;
for fraction in fraction_correct:
    a = a + fraction;
average_fraction_correct_forecast_C = a / 100
print("%.3f" % average_fraction_correct_forecast_C)
# YOUR CODE HERE
```

<!-- grade_id: cell-90b8668d530f9ac2, points: 0, grade: false, solution: true -->
```python
expert_averages = {"First": "A", "Second": "C", "Third": "B"}

# YOUR CODE HERE
```

<!-- grade_id: cell-c66887f6e48fab6d, points: 1, grade: true, solution: false -->
```python
# TEST 

assert "A" in expert_averages.values()
assert "B" in expert_averages.values()
assert "C" in expert_averages.values()

```

#### Question 1.5

Having evaluated experts based on average fraction of correct forecasts in different market conditions, we want to evaluate experts on consistency. We measure consistency as the longest stretch of consecutive correct forecasts.

```python
expert = "C"
number_forecasts = 100
market_conditions = 1

forecasts = generate_expert_prediction(expert, number_forecasts, market_conditions)
print(forecasts)
```

Here we obtained `100` forecasts in market conditions `1` from expert `C`. 

We can use the function `get_consecutive` in the package `helper_functions` to calculate the length of the longest strech of consecutive correct forecasts.

```python
number_consecutive_correct = get_consecutive(forecasts)
print(number_consecutive_correct)
```

Here we learn that expert `C` made 14 consecutive correct forecasts over the course of `100` forecasts in market conditions `1`.

Repeat the calculation with expert `B`. 

> How long was longest stretch of consecutive correct forecasts made by expert `B` over the course of `100` forecasts in market conditions `1`? 

Enter your answer as an integer number to the variable `q15`

```python
expert = "B"
number_forecasts = 100
market_conditions = 1

forecasts = generate_expert_prediction(expert, number_forecasts, market_conditions)
print(forecasts)
number_consecutive_correct = get_consecutive(forecasts)
print(number_consecutive_correct)
```

<!-- grade_id: cell-e10fd7818db33b12, points: 0, grade: false, solution: true -->
```python
q15 = 5

# YOUR CODE HERE
```

<!-- grade_id: cell-f0a7f3f99fa2e118, points: 1, grade: true, solution: false -->
```python
# TEST 

assert 0 < q15 < 10

```

#### Question 1.6

Suppose we want to repeat the calculation from Question 1.5 in different market conditions. We will study market conditions

> 1,2,3,...,98,99,100

We will allow expert `C` to make 1000 forecasts in each of the 100 market conditions. We want to calculate the length of the longest stretch of consecutive correct forecasts.

Complete the following code by filling in the blank.

<!-- grade_id: cell-2b00908b752c4c7c, points: 0, grade: false, solution: true -->
```python
expert = "C"
number_forecasts = 1000
market_conditions = range(1,101)

consecutive_correct_forecasts = []
for market_condition in market_conditions:
    forecasts = generate_expert_prediction(expert, number_forecasts, market_condition)
    
    # YOUR CODE HERE
    
    consecutive_correct_forecasts.append(get_consecutive(forecasts))
```

We obtain 100 numbers corresponding to the 100 market conditions. 

<!-- grade_id: cell-f052d5619ddc859b, points: 0, grade: true, solution: false -->
```python
print(consecutive_correct_forecasts)
```

Compute the average of these numbers. Note that your answer must be a floating point number correct to three decimal places. 

<!-- grade_id: cell-e76326379f03ee88, points: 0, grade: false, solution: true -->
```python
a = 0;
for foreast in consecutive_correct_forecasts:
    a = a + foreast;
consecutive_correct_forecast_C = a / 100
print(consecutive_correct_forecast_C)
# YOUR CODE HERE
```

<!-- grade_id: cell-484994a47c8bf17b, points: 1, grade: true, solution: false -->
```python
# TEST 

assert 0 < consecutive_correct_forecast_C < 50

```

```python
consecutive_correct_forecast_C
```

#### Question 1.7

Repeat the calculation from Question 1.6 for expert `A` and `B`. Based on the average length of the longest stretch of consecutive correct forecasts, we can rank the experts 

- `First`
- `Second`
- `Third`

Note that `First` corresponds to the longest average  and `Third` corresponds to the shortest average. Fill in the keys of the dictionary `expert_consecutive` with the string 

- `A`
- `B`
- `C`

```python
expert = "A"
number_forecasts = 1000
market_conditions = range(1,101)

consecutive_correct_forecasts = []
for market_condition in market_conditions:
    forecasts = generate_expert_prediction(expert, number_forecasts, market_condition)
    
    # YOUR CODE HERE
    
    consecutive_correct_forecasts.append(get_consecutive(forecasts))
    
a = 0;
for foreast in consecutive_correct_forecasts:
    a = a + foreast;
consecutive_correct_forecast_A = a / 100
print(consecutive_correct_forecast_A)
# YOUR CODE HERE
```

```python
expert = "B"
number_forecasts = 1000
market_conditions = range(1,101)

consecutive_correct_forecasts = []
for market_condition in market_conditions:
    forecasts = generate_expert_prediction(expert, number_forecasts, market_condition)
    
    # YOUR CODE HERE
    
    consecutive_correct_forecasts.append(get_consecutive(forecasts))
    
a = 0;
for foreast in consecutive_correct_forecasts:
    a = a + foreast;
consecutive_correct_forecast_B = a / 100
print(consecutive_correct_forecast_B)
# YOUR CODE HERE
```

<!-- grade_id: cell-ab8a75f39c805703, points: 0, grade: false, solution: true -->
```python
expert_consecutive = {"First": "C", "Second":"A", "Third": "B"}

# YOUR CODE HERE
```

<!-- grade_id: cell-e9e81e67e8cef318, points: 1, grade: true, solution: false -->
```python
# TEST 

assert "A" in expert_consecutive.values()
assert "B" in expert_consecutive.values()
assert "C" in expert_consecutive.values()

```

### 2. Packages 

The Federal Reserve Bank of St. Louis publishes data about jobs in the US.  We have data from a [government website](https://fred.stlouisfed.org/categories/33509) stored in the file `unemployment.csv`. There are many ways of defining unemployment. The dataset includes two notions of the unemployment rate:

1. Among people who are able to work and are looking for a full-time job, the percentage who can't find a job.  This is called the Non-Employment Index, or NEI.
2. Among people who are able to work and are looking for a full-time job, the percentage who can't find any job *or* are only working at a part-time job.  The latter group is called "Part-Time for Economic Reasons", so the acronym for this index is NEI-PTER.  

#### Question 2.1 

Import the function `load_data` from the module `helper_functions` to load the data into a variable called `unemployment`. Note that the input is the string `"unemployment.csv"`.

<!-- grade_id: cell-c2c07f92b6e89a1b, points: 0, grade: false, solution: true -->
```python
from helper_functions import load_data
unemployment = load_data("unemployment.csv")
print(unemployment)

# YOUR CODE HERE
```

<!-- grade_id: cell-7f1d6aef2a21a43a, points: 1, grade: true, solution: false -->
```python
# TEST 

assert len(unemployment) == 90

```

#### Question 2.2

Note that we have three columns `Date`, `NEI` and `NEI-PTER`. 

```python
unemployment
```

Based on the entries in the `Date` columns, does the Federal Reserve Bank report the data

- annually 
- quarterly 
- monthly 

<!-- grade_id: cell-0127e6ab02022628, points: 0, grade: false, solution: true -->
```python
q22 = "quarterly"

# YOUR CODE HERE
```

<!-- grade_id: cell-abff9489003728c0, points: 1, grade: true, solution: false -->
```python
# TEST 

assert q22 in ["quarterly", "anually", "monthly"]

```

#### Question 2.3
Note that NEI-PTER is the percentage of people who are unemployed (and counted in the NEI) plus the percentage of people who are PTER.  

Import the function `add_pter` from the module `helper_functions` to compute the PTER from the NEI-PTER and NEI. The input is `unemployment` from Question 2.1. You should assign the output to a variable called `unemployment_with_pter`


<!-- grade_id: cell-4d820c0947ceca63, points: 0, grade: false, solution: true -->
```python
from helper_functions import add_pter
unemployment_with_pter = add_pter(unemployment)
print(unemployment_with_pter)
# YOUR CODE HERE
```

<!-- grade_id: cell-336a9ce086d3dbdd, points: 1, grade: true, solution: false -->
```python
# TEST 

assert len(unemployment_with_pter["PTER"]) == 90
```

#### Question 2.4
Many economists feel that PTER increased in the "Great Recession" of 2008-2009.

Import the function `plot_pter` from the module `helper_functions` to generate a plot of the PTER. The input is `unemployment_with_pter` from Question 2.3.

<!-- grade_id: cell-b772a16f4ab3b375, points: 0, grade: false, solution: true -->
```python
from helper_functions import plot_pter
plot_pter(unemployment_with_pter)
# YOUR CODE HERE
```

True or False : Based on the chart, were PTER rates high during or directly after the Great Recession?

<!-- grade_id: cell-93c470d443512cce, points: 0, grade: false, solution: true -->
```python
q23 = True

# YOUR CODE HERE
```

<!-- grade_id: cell-1ec6812427a8921d, points: 1, grade: true, solution: false -->
```python
# TEST 

assert q23 in [True, False]

```

**Dumped cells**

```python
expert = "B"
number_forecasts = 1000
market_conditions = range(1,101)
```

```python
forecasts = generate_expert_prediction(expert, number_forecasts, market_conditions[0])

number_correct = 0 
for forecast in forecasts:
    if (forecast == 1):
        number_correct = number_correct + 1
        
print(number_correct / number_forecasts)
```

```python
forecasts = generate_expert_prediction(expert, number_forecasts, market_conditions[0])

number_correct = 0 
for forecast in forecasts:
    if (forecast == 1):
        number_correct = number_correct + 1
        
print(number_correct / number_forecasts)
```

```python
fraction_correct = []
for market_condition in market_conditions:
    forecasts = generate_expert_prediction(expert, number_forecasts, market_condition)

    number_correct = 0 
    for forecast in forecasts:
        if (forecast == 1):
            number_correct = number_correct + 1
        
    fraction_correct.append(number_correct/number_forecasts)
    
print(fraction_correct)
```

<!-- grade_id: cell-08f7c5b916ea3087, points: 0, grade: false, solution: true -->
```python
a = 0;
for fraction in fraction_correct:
    a = a + fraction;
average_fraction_correct_forecast_B = a / 100
# YOUR CODE HERE

```